```python
# Initialize Otter
import otter
grader = otter.Notebook("demo_master.ipynb")
```

```python
import math
```

**Question 1**. Write a function which creates a rectangular triangle given hypotenuse and one cathetus. Check if the one could be created from the given inputs by using Pythagorean theorem to compute the second cathetus and making sure it is a real number.

```python
def create_triangle(cathetus1, hypotenuse):
    '''
    Create a rectangular triangle as a string of stars
    '''
    ...
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 2)
```
>>> def test_triangles(create_triangle):
...     assert create_triangle(3, 5) == '*\n**\n***'
...     assert create_triangle(6, 10) == '*\n**\n***\n****\n*****\n******'
...     try:
...         create_triangle(3, 4)
...     except ValueError as e:
...         assert str(e) == 'Invalid right triangle dimensions'
>>> test_triangles(create_triangle)
```
